# [ERG-190C] Homework 7: Gradient Descent

In this homework, we explore modeling data, estimating optimal parameters and a numerical estimation method, gradient descent. These concepts are some of the fundamentals of data science and machine learning and will serve as the building blocks for future projects, classes, and work.

After this homework, you should feel comfortable with the following:

+ Practice reasoning about a model

+ Build some intuition for loss functions and how the behave 

+ Work through deriving the gradient of a loss with respect to model parameters

+ Work through a basic version of gradient descent.

This homework is comprised of completing code, deriving analytic solutions, writing LaTex and visualizing loss.

### Table of Contents
1 - [A Simple Model](#section model)<br>
2 - [Fitting the Model](#section fitting)<br>
3 - [Increasing Model Complexity](# section complexity)<br>
4 - [Gradient Descent](#section gd)<br>
5 - [Bonus: Visualizing Loss](#section loss)<br>

<br>
**Dependencies:**

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import csv
import warnings 
warnings.filterwarnings('ignore')
plt.style.use('fivethirtyeight') 

# Set some parameters
plt.rcParams['figure.figsize'] = (12, 9)
plt.rcParams['font.size'] = 16
np.set_printoptions(4)
```

```python
# We will use plot_3d helper function to help us visualize gradient
from hw7_utils import plot_3d
```

----

## Load Data
Load the data.csv file into a pandas dataframe.  
Note that we are reading the data directly from the URL address.

```python
# Run this cell to load our sample data
data = pd.read_csv("http://www.ds100.org/sp18/assets/datasets/hw5_data.csv", index_col=0)
data.head()
```

---

## Section 1. A Simple Model<a id='section model'></a>
Let's start by examining our data and creating a simple model that can represent this data.

<br>**Question 1.1:** First, let's visualize the data in a scatter plot. After implementing the `scatter` function below, you should see something like this:
![scatter](scatter.png)


```python
def scatter(x, y):
    """
    Generate a scatter plot using x and y

    Keyword arguments:
    x -- the vector of values x
    y -- the vector of values y
    """
    plt.figure(figsize=(8, 6))
    ...
    # YOUR CODE HERE

x = data['x']
y = data['y']
scatter(x,y)
plt.show()
```

**Question 1.2:** Describe any significant observations about the distribution of the data. How can you describe the relationship between $x$ and $y$?

<!-- grade_id: q1b-answer, points: 1, grade: true, solution: true -->
Answer: ...

**Question 1.3:** The data looks roughly linear, with some extra noise. For now, let's assume that the data follows some linear model, parametrized by $\theta$:

$$\Large
\hat{y} = \theta \cdot x
$$

Define a linear model function that estimates a value $\hat{y}$ given $x$ and $\theta$.


<!-- grade_id: q1c-answer, points: 0, grade: false, solution: true -->
```python
def linear_model(x, theta):
    """
    Returns the estimate of y given x and theta

    Keyword arguments:
    x -- the vector of values x
    theta -- the scalar theta
    """
    y_hat = ...
    # YOUR CODE HERE
    return y_hat
```

```python
# run this cell
linear_model(0, 1)
```

```python
# run this cell
np.sum(linear_model(np.array([3, 5]), 3))
```

```python
# run this cell
linear_model(np.array([7, 8]), 4).mean()
```

**Question 1.4:** In class, we learned that the $L^2$ (or squared) loss function is smooth and continuous. Let's use $L^2$ loss to evaluate our estimate on $\theta$ and predict an optimal value for $\theta$, named $\theta^*$. Define the $L^2$ loss function `l2_loss` below.

<!-- grade_id: q1d-answer, points: 0, grade: false, solution: true -->
```python
def l2_loss(y, y_hat):
    """
    Returns the average l2 loss given y and y_hat

    Keyword arguments:
    y -- the vector of true values y
    y_hat -- the vector of predicted values y_hat
    """
    ...
    # YOUR CODE HERE
```

```python
# run this cell
l2_loss(2, 1)
```

```python
# run this cell
l2_loss(np.array([1, 1, 1]), np.array([4, 1, 4]))
```

**Question 1.5:** First, visualize the: $L^2$ loss as a function of $\theta$, where several different values of $\theta$ are given. Be sure to label your axes properly. You plot should look something like this:
![avg_l2](l2_avg_loss.png)

What looks like the optimal $\theta^*$ value based on the visualization? Set `theta_star_guess` to the value of $\theta$ that appears to minimize our loss.

<!-- grade_id: q1e-answer, points: 0, grade: false, solution: true -->
```python
def visualize(x, y, thetas):
    """
    Plots the average l2 loss for given x, y as a function of theta.
    Use the functions you wrote for linear_model and l2_loss.

    Keyword arguments:
    x -- the vector of values x
    y -- the vector of values y
    thetas -- the vector containing different estimates of theta
    """
    y_hat = linear_model(x, theta)
    
    avg_loss = ... # Calculate the loss here for each value of theta
    
    plt.figure(figsize=(8,6))
    
    ... # Create your plot here
    
    # YOUR CODE HERE
    
thetas = np.linspace(-1, 5, 70)
visualize(x, y, thetas)
plt.show()
theta_star_guess = ...
# YOUR CODE HERE
```

```python
print(l2_loss(3, 2) == 1)
print(l2_loss(0, 10) == 100)
print(1 <= theta_star_guess <= 2)
```

---
## Section 2: Fitting our Simple Model<a id='section fitting'></a>
Now that we have defined a simple linear model and loss function, let's begin working on fitting our model to the data.

**Question 2.1:** Let's confirm our visual findings for optimal $\theta^*$. First, find the analytical solution for the optimal $\theta^*$ for average $L^2$ loss. Of the three options, below, which is the correct answer? Highlight your answer in <font color = "red">red</font> (double click this cell if you don't know how to do this).

1. $$\Large {\theta}^* = \frac{\sum x_i + y_i}{\sum x_i^2}$$ <br>
2. $$\Large {\theta}^* = \frac{\sum x_iy_i}{\sum x_i}$$ <br>
3. $$\Large {\theta}^* = \frac{\sum x_iy_i}{\sum x_i^2}$$ <br>


**Question 2.2:** 
Now that we have the analytic solution for $\theta^*$, implement the function `find_theta` that calculates the numerical value of $\theta^*$ based on our data $x$, $y$.



<!-- grade_id: q2b-answer, points: 0, grade: false, solution: true -->
```python
def find_theta(x, y):
    """
    Find optimal theta given x and y

    Keyword arguments:
    x -- the vector of values x
    y -- the vector of values y
    """
    theta_opt = ...
    # YOUR CODE HERE
    return theta_opt
```

<!-- grade_id: q2b-tests, points: 1, grade: true, solution: false -->
```python
t_star = find_theta(x, y)
print(f'theta_opt = {t_star}')

assert 1.4 <= t_star <= 1.6
```

**Question 2.3:** Now, let's plot our loss function again using the `visualize` function. But this time, add a vertical line at the optimal value of theta (plot the line $x = \theta^*$). Your plot should look something like this:
![vertical_linear](vertical_linear.png)

<!-- grade_id: q2c-answer, points: 1, grade: true, solution: true -->
```python
theta_opt = ...
...
...
# YOUR CODE HERE.  Hint to simplify things: there is a special matplotlib function for plotting vertical lines.  
plt.show()
```

<br> 
**Question 2.4:** We now have an optimal value for $\theta$ that minimizes our loss. In the cell below, plot the scatter plot of the data from Question 1a (you can reuse the `scatter` function here). But this time, add the line $\hat{y} = \theta^* \cdot x$ using the $\theta^*$ you computed above. Your plot should look something like this:
![scatter_with_line](scatter_with_line.png)

<!-- grade_id: q2d-answer, points: 1, grade: true, solution: true -->
```python
theta_opt = ...
...
...
# YOUR CODE HERE
plt.show()
```

<br>
**Question 2.5:** Great! It looks like our estimate for $\theta$ is able to capture a lot of the data with a single parameter. Now let's try to remove the linear portion of our model from the data to see if we missed anything. 

The remaining data is known as the residual, $r=y-\theta^* \cdot x$. Below, write a function to find the residual and plot the residuals corresponding to x in a scatter plot. Plot a horizontal line at $y=0$ to assist visualization.

<!-- grade_id: q2e-answer, points: 1, grade: true, solution: true -->
```python
def visualize_residual(x, y):
    """
    Plot a scatter plot of the residuals, the remaining 
    values after removing the linear model from our data.

    Keyword arguments:
    x -- the vector of values x
    y -- the vector of values y
    """
    ...
    # YOUR CODE HERE

visualize_residual(x, y)
```

**Question 2.6:** What does the residual look like? Do you notice a relationship between $x$ and $r$?

<!-- grade_id: q2f-answer, points: 1, grade: true, solution: true -->
Answer: ...

---
## Section 3: Increasing Model Complexity<a id='section complexity'></a>

It looks like the remaining data is sinusoidal, meaning our original data follows a linear function and a sinusoidal function. Let's define a new model to address this discovery and find optimal parameters to best fit the data:

$$\Large
\hat{y} = \theta_1x + sin(\theta_2x)
$$

Now, our model is parameterized by both $\theta_1$ and $\theta_2$, or composed together, $\vec{\theta}$.

Note that a generalized sine function $a\sin(bx+c)$ has three parameters: amplitude scaling parameter $a$, frequency parameter $b$ and phase shifting parameter $c$. We can assume that the scaling and shifting parameter ($a$ and $c$ in this case) are 1 and 0 respectively. 

<br>
**Question 3.1:** In the following cell, **explain why we can assume the scaling parameter to be 1 and shifting parameter to be 0 based on the residual plot in Question 2e**. 

You might find the following code helpful in visualizing all three parameters.

```python
def plot_sin_generalized(a,b,c,label=None):
    """Plot a sin function with three parameters"""
    X = np.linspace(-5, 5)
    Y = a * np.sin(b*X + c)
    plt.scatter(X, Y, label=label)
    plt.legend()
```

You can try plotting: 
```python
plot_sin_generalized(1,1,1, label='sin(x)')
plot_sin_generalized(1,1,2, label='sin(x + 2)')
plot_sin_generalized(1,2,2, label='sin(2x + 2)')
plot_sin_generalized(2,2,2, label='2sin(2x + 2)')
```

<!-- grade_id: q3a-answer, points: 1, grade: true, solution: true -->
Answer: ***YOUR ANSWER HERE***

Possible Answer: The shifting parameter is 0 because the graph has a point at the origin, meaning it is not shifted left or right.
The scaling parameter is 1 because the points seem to be spread around a sine wave going from -1 to 1. 

**Question 3.2:** As in Question 1, write a function that predicts a value $\hat{y}$ given an input $x$ based on our new model.

*Hint:* Try to do this without using for loops. The `np.sin` function may help you.

<!-- grade_id: q3b-answer, points: 0, grade: false, solution: true -->
```python
def sin_model(x, theta_1, theta_2):
    """
    Predict the estimate of y given x, theta_1, theta_2

    Keyword arguments:
    x -- the vector of values x
    theta_1 -- the scalar value theta_1
    theta_2 -- the scalar value theta_2
    """
    y_hat = ...
    # YOUR CODE HERE
    return y_hat
```

<!-- grade_id: q3b-tests, points: 1, grade: true, solution: false -->
```python
print(np.isclose(sin_model(1, 1, np.pi), 1.0000000000000002))
# Check that we accept x as arrays
assert len(sin_model(x, 2, 2)) > 1
```

**Question 3.3:** In this question your job is to match the left and right sides of the equations for:
1. The $L^2$ loss for for the `sin` model, $\hat{y} = \theta_1x + sin(\theta_2x)$.  We'll call that $L(x, y, \theta_1, \theta_2)$.
2. The partial derivatives of the `sin` model loss functions, $\frac{\partial L }{\partial \theta_1}, \frac{\partial L }{\partial \theta_2}$. 

Notice that we now have $\vec{x}$ and $\vec{y}$ instead of $x$ and $y$. This means that when determining the loss function $L(x, y, \theta_1, \theta_2)$, you'll need to take the average of the squared losses for each $y_i$, $\hat{y_i}$ pair.

As your answer below, match the right side (letters) to the correct left sides (numbers).

<!-- grade_id: q3c-answer, points: 1, grade: true, solution: true -->


1. $L(x, y, \theta_1, \theta_2)$ <br>
2. $\frac{\partial L}{\partial \theta_1}$ <br>
3. $\frac{\partial L}{\partial \theta_2}$ <br>



A.  $\frac{1}{n} \sum_{i=1}^n (y_i - \theta_1 x_i - \sin(\theta_2 x_i)) ^ 2$ <br>
B. $-\frac{2}{n} \sum_{i=1}^n (x_i y_i \cos(\theta_2 x_i) - \theta_1 x_i ^ 2 \cos(\theta_2 x_i) - x_i \sin(\theta_2 x_i)\cos(\theta_2 x_i))$ <br>
C. $-\frac{2}{n} \sum_{i=1}^n (x_i y_i - \theta_1 x_i ^ 2 - x_i \sin(\theta_2 x_i))$


Answer: <br>
1. ... <br>
2. ... <br>
3. ... <br>

**Question 3.4:** Now, implement the functions `dt1` and `dt2`, which should compute $\frac{\partial L }{\partial \theta_1}$ and $\frac{\partial L }{\partial \theta_2}$ respectively. Use the formulas you wrote for $\frac{\partial L }{\partial \theta_1}$ and $\frac{\partial L }{\partial \theta_2}$ in the previous exercise. In the functions below, the parameter `theta` is a vector that looks like $( \theta_1, \theta_2 )$.

<!-- grade_id: q3d-answer-1, points: 0, grade: false, solution: true -->
```python
def dt1(x, y, theta):
    """
    Compute the numerical value of the partial of l2 loss with respect to theta_1

    Keyword arguments:
    x -- the vector of all x values
    y -- the vector of all y values
    theta -- the vector of values theta
    """
    ...
    # YOUR CODE HERE
```

<!-- grade_id: q3d-answer-2, points: 0, grade: false, solution: true -->
```python
def dt2(x, y, theta):
    """
    Compute the numerical value of the partial of l2 loss with respect to theta_2

    Keyword arguments:
    x -- the vector of all x values
    y -- the vector of all y values
    theta -- the vector of values theta
    """
    ...
    # YOUR CODE HERE
```

```python
# This function calls dt1 and dt2 and returns the gradient dt. It is already implemented for you.
def dt(x, y, theta):
    """
    Returns the gradient of l2 loss with respect to vector theta

    Keyword arguments:
    x -- the vector of values x
    y -- the vector of values y
    theta -- the vector of values theta
    """
    return np.vstack([
        dt1(x, y, theta),
        dt2(x, y, theta)
    ])
```

<!-- grade_id: q3d-tests, points: 1, grade: true, solution: false -->
```python
print(np.isclose(dt1(x, y, [0, np.pi]), -25.376660670924529))
print(np.isclose(dt2(x, y, [0, np.pi]), 1.9427210155296564))
```

---
## Section 4: Gradient Descent<a id='section gd'></a>
Now try to solve for the optimal $\theta^*$ analytically...

**Just kidding!**

You can try but we don't recommend it. When finding an analytic solution becomes difficult or impossible, we resort to alternative optimization methods for finding an approximate solution.

So let's try implementing a numerical optimization method: gradient descent!

<br>
**Question 4.1:** Implement the `grad_desc` function that performs gradient descent for a finite number of iterations. This function takes in array $x$, array $y$, and an initial value for $\theta$ (`theta`). `alpha` will be the learning rate (or step size, whichever term you prefer). In this part, we'll use a static learning rate that is the same at every time step. 

At each time step, use the gradient and `alpha` to update your current `theta`. Also at each time step, be sure to save the current `theta` in `theta_history`, along with the $L^2$ loss (computed with the current `theta`) in `loss_history`.

Hints:
- Write out the gradient update equation (1 step). What variables will you need for each gradient update? Of these variables, which ones do you already have, and which ones will you need to recompute at each time step?
- You may need a loop here to update `theta` several times.

```python
# Run me
def init_t():
    """Creates an initial theta [0, 0] as a starting point for gradient descent"""
    return np.zeros((2,1))
```

<!-- grade_id: q4a-answer, points: 0, grade: false, solution: true -->
```python
def grad_desc(x, y, theta, num_iter=20, alpha=0.1):
    """
    Run gradient descent update for a finite number of iterations and static learning rate

    Keyword arguments:
    x -- the vector of values x
    y -- the vector of values y
    theta -- the vector of values theta to use at first iteration
    num_iter -- the max number of iterations
    alpha -- the learning rate (also called the step size)
    
    Return:
    theta -- the optimal value of theta after num_iter of gradient descent
    theta_history -- the series of theta values over each iteration of gradient descent
    loss_history -- the series of loss values over each iteration of gradient descent
    """
    theta_history = []
    loss_history = []
    
    ...
    
    # YOUR CODE HERE
    return theta, theta_history, loss_history
```

<!-- grade_id: q4a-tests, points: 1, grade: true, solution: false -->
```python
t = init_t()
t_est, ts, loss = grad_desc(x, y, t, num_iter=20, alpha=0.1)

print(len(ts) == len(loss) == 20) # theta history and loss history are 20 items in them
print(ts[0].shape == (2,1)) # theta history contains theta values
print(np.isscalar(loss[0])) # loss history is a list of scalar values, not vector

print(loss[1] - loss[-1] > 0) # loss is decreasing

print(np.allclose(np.sum(t_est), 4.5, atol=2e-1))  # theta_est should be close to our value
```

**Question 4.2:** Let's visually inspect our results of running gradient descent to optimize $\theta$. Plot our x values with our model's predicted y values over the original scatter plot. Did gradient descent successfully optimize $\theta$?

```python
# Run me
t = init_t()
t_est, ts, loss = grad_desc(x, y, t)
```

```python
y_pred = sin_model(x, t_est[0], t_est[1])

plt.plot(x, y_pred, label='Model')
plt.scatter(x, y, alpha=0.5, label='Observation', color='gold')
plt.legend()
plt.show()
```

**Question 4.3:** Let's visualize gradient descent to see how it converges. Plot the loss values over each iteration of gradient descent. 

<!-- grade_id: q4d-answer, points: 1, grade: true, solution: true -->
```python
# plt.plot(...) # Plot of loss history for static learning rate
...
# YOUR CODE HERE
```

```python
plt.plot(loss)
plt.xlabel('Iteration')
plt.ylabel('Loss Value')
plt.title("Learning Rate")
plt.show()
```

**Question 4.4:** Describe the approach of modifying the alpha parameter on each iteration and how a plot of loss versus iteration number would differ from the one you just made. Give some intuition for your answer.

<!-- grade_id: q4e-answer, points: 1, grade: true, solution: true -->
Answer: ...

----

## Submission

Congrats, you're done with homework 7! 

In order to turn in this assignment, go to the toolbar and click **File** -> **Download as** -> **.html** and submit the file through bCourses.

If you want, you can move the bonus section down below!

---

## Section 5: Visualizing Loss [Bonus]<a id='section loss'></a>

**Question 5.1 - 3D Plot:**
In the previous plot is about the loss decrease over time, but what exactly is path the theta value? Run the following three cells. 

```python
# Run me
#ts = np.array(ts).squeeze()
#ts_decay = np.array(ts_decay).squeeze()
#loss = np.array(loss)
#loss_decay = np.array(loss_decay)
```

```python
# Run me to see a 3D plot (gradient descent with static alpha)
plot_3d(ts[:, 0], ts[:, 1], loss, l2_loss, sin_model, x, y)
```

```python
# Run me to see another 3D plot (gradient descent with decaying alpha)
plot_3d(ts_decay[:, 0], ts_decay[:, 1], loss_decay, l2_loss, sin_model, x, y)
```

In the following cell, write 1-2 sentences about the differences between using a static learning rate and a learning rate with decay for gradient descent. Use the loss history plot as well as the two 3D visualization to support your answer.

<!-- grade_id: q5a-answer, points: 1, grade: true, solution: true -->
Answer: ...

** Question 5.2 - Contour Plot:**

Another common way of visualizing 3D dynamics is with a _contour_ plot. 

Please refer to this notebook when you are working on the next question: http://www.ds100.org/fa17/assets/notebooks/26-lec/Logistic_Regression_Part_2.html#Visualizing-the-gradient-descent-path

In next question, fill in the necessary part to create a contour plot. Then run the following cells. 

```python
## Run me
import plotly
import plotly.graph_objs as go
plotly.offline.init_notebook_mode(connected=True)
```

<!-- grade_id: q5b-1, points: 0, grade: false, solution: true -->
```python
def contour_plot(title, theta_history, loss_function, model, x, y):
    """
    The function takes the following as argument:
        theta_history: a (N, 2) array of theta history
        loss: a list or array of loss value
        loss_function: for example, l2_loss
        model: for example, sin_model
        x: the original x input
        y: the original y output
    """
    theta_1_series = theta_history[:,0] # a list or array of theta_1 value
    theta_2_series = theta_history[:,1] # a list or array of theta_2 value
    
    # Create trace of theta point
    thata_points = go.Scatter(name="Theta Values", 
                              x=..., #TODO
                              y=..., #TODO
                              mode="lines+markers")

    ## In the following block of code, we generate the z value
    ## across a 2D grid
    t1_s = np.linspace(np.min(theta_1_series) - 0.1, np.max(theta_1_series) + 0.1)
    t2_s = np.linspace(np.min(theta_2_series) - 0.1, np.max(theta_2_series) + 0.1)

    x_s, y_s = np.meshgrid(t1_s, t2_s)
    data = np.stack([x_s.flatten(), y_s.flatten()]).T
    ls = []
    for t1, t2 in data:
        l = loss_function(model(x, t1, t2), y)
        ls.append(l)
    z = np.array(ls).reshape(50, 50)
    
    # Create the contour 
    lr_loss_contours = go.Contour(x=..., #TODO 
                                  y=..., #TODO
                                  z=..., #TODO
                                  colorscale='Viridis', reversescale=True)
    
    
    # YOUR CODE HERE

    plotly.offline.iplot(go.Figure(data=[lr_loss_contours, thata_points], layout={'title': title}))
```

```python
# Run this
#contour_plot('Gradient Descent with Static Learning Rate', ts, l2_loss, sin_model, x, y)
```

```python
## Run me
#contour_plot('Gradient Descent with Decay Learning Rate', ts_decay, l2_loss, sin_model, x, y)
```

In the following cells, write down the answer to the following questions:
- How do you interpret the two contour plots? 
- Compare contour plot and 3D plot, what are the pros and cons of each? 

<!-- grade_id: q5b-5, points: 1, grade: true, solution: true -->
Answer: ...

### How to Improve?
It looks like our model, a combination of a linear function and sinusoidal function, was able to almost perfectly fit our data. It turns out that many real world scenarios come from relatively simple models. At the same time, the real world can be incredibly complex and a simple model wouldn't work so well. Consider the example below; it is neither linear, nor sinusoidal, nor quadratic. Suggest how we could iteratively create a model to fit this data and how we might improve our results. Try fitting the data below if you have time :-)

```python
x = []
y = []
for t in np.linspace(0,10*np.pi, 200):
    r = ((t)**2)
    x.append(r*np.cos(t))
    y.append(r*np.sin(t))

plt.scatter(x,y)  
plt.show()
```

<!-- grade_id: improvements-suggestion, points: 1, grade: true, solution: true -->
Answer: ...


----

## Bibliography

+ Data 100 - HW 5: Modeling, Estimation and Gradient Descent

<hr/>

Data Science Modules: http://data.berkeley.edu/education/modules